```python
# Initialize Otter
import otter
grader = otter.Notebook("01_dataset_manipulation.ipynb")
```

# Assignment: Data cleaning, reshaping, merging and probability

**Author**: (Write your name here)

## Introduction

This assignment is meant to give you an opportunity to use the tools that we've discussed for the last two weeks. We will be downloading data from the World Bank _World Development Indicators_ and doing a "guided case-study."

In order to run the code below, you will need to install the `world_bank_data` package. This package has built a client library for the World Bank Data API which saves us time by allowing us simplified methods for interfacing with the data. the package can be installed by uncommenting the cell below and running it. 

## Section 1: Data ingestion


```python
# if the next cell fails, uncomment the line below and re-run
# %pip install world_bank_data
```

```python
import matplotlib.pyplot as plt
import numpy as np
import pandas as pd

import world_bank_data as wb

%matplotlib inline
```

In this section, we are going to read in a few different datasets from the World Bank.

* **Primary completion rate**, or gross intake ratio to the last grade of primary education, is the number of new entrants (enrollments minus repeaters) in the last grade of primary education, regardless of age, divided by the population at the entrance age for the last grade of primary education. Data limitations preclude adjusting for students who drop out during the final year of primary education.
* **GDP** at purchaser's prices is the sum of gross value added by all resident producers in the economy plus any product taxes and minus any subsidies not included in the value of the products. It is calculated without making deductions for depreciation of fabricated assets or for depletion and degradation of natural resources. Data are in current U.S. dollars.
* **Population** is based on the de facto definition of population, which counts all residents regardless of legal status or citizenship. The values shown are midyear estimates.

```python
countries = wb.get_countries()

dates = "2000:2019"
educ_primary_completion = wb.get_series("SE.PRM.CMPT.ZS", date=dates)
gdp = wb.get_series("NY.GDP.MKTP.CD", date=dates)
populations = wb.get_series("SP.POP.TOTL", date=dates)
```

```python
countries.head()
```

```python
educ_primary_completion.head()
```

```python
gdp.head()
```

```python
populations.head()
```

```python
countries.head()
```

### Question 1:

We want to create a dataset that has `(country, year)` on the index with  `gdp` (GDP), `pop` (total population), , and `pcr` (primary completion rate) as the columns. We only want to include countries in this dataset so you will need to drop any of the aggregated regions reported by the World Bank (i.e. there are reported values for "Africa" and "Asia" but these are aggregated numbers rather than country specific so we would drop these and the other aggregates -- Hint: the `countries` DataFrame should help you with this).

Use the data that we read in above to create this dataset and call the new dataset `df_full`.

After you've created the full dataset, create a new dataset with only the observations from the year 2018. Call this dataset `df` -- We will use this subset of data for the remainder of the assignment.

```python
...
```

```python
grader.check("q1")
```

### Question 2:

Create a new variable, `loggdp_percapita` that defined as $\log \left( \frac{\text{GDP}}{\text{population}} \right)$.

```python
...
```

```python
grader.check("q2")
```

<!-- BEGIN QUESTION -->

### Question 3:

Create a scatter plot with `loggdp_percapita` on the x-axis and `pcr` on the y-axis.

What relationships do you observe between these two variables?

```python
...
```

<!-- END QUESTION -->

### Question 4:

Read the documentation for the `np.histogram2d` function. Create a two-dimensional histogram and store it as `count` and the bin values as `xvals` and `yvals`. Set the `x` to `loggdp_percapita` and `y` to `pcr`. Let there be 8 bins the `x` direction and 10 in the `y` direction.

Note that you need to remove missing values before using `np.histogram2d` -- we've done this for you in the variable `df_no_nan`

```python
df_no_nan = df.dropna()

...
```

```python
grader.check("q4")
```

### Question 5:

Turn the numpy array `count` into a joint probability distribution called `jd`.

```python
jd = ...
```

```python
grader.check("q5")
```

### Question 6:

Compute the marginal distributions for `loggdp_percapita` and `pcr`

Call them `marginal_lgdp` and `marginal_pcr`

```python
marginal_lgdp = ...
marginal_pcr = ...
marginal_pcr
```

```python
grader.check("q6")
```

### Question 7:

For each "primary schooling completed" bin, compute the conditional distribution across log GDP per capita.

```python
conditionals_lgdp = ...

conditionals_lgdp
```

```python
grader.check("q7")
```

<!-- BEGIN QUESTION -->

### Question 8:

Is the joint distribution created from the histogram the only joint distribution that could match with the marginal distributions found in question 6? If not, please find one more joint distribution that could describe the marginals.

_Type your answer here, replacing this text._

```python
...
```

<!-- END QUESTION -->

### Question 9:

Consider daily sales data from a local convenience store. The data is broken into 12 datasets (one for each month). What function do you think is best to combine this data into a single dataset? Why?

Below is a fake dataset with this information. Please use the function you mention previously to create a `df_sales` using the `jan_sales`, `feb_sales`, ..., `dec_sales` DataFrames.

Answer here about which function and why

```python
# Don't change this cell -- It creates the dataset for you
import calendar

weeks = pd.date_range("2020-01-01", "2020-12-31", freq="W")

dfs_sales = []
for month in range(1, 13):
    days = calendar.monthrange(2020, month)
    weeks = pd.date_range(
        f"2020-{month:0>2}-01", f"2020-{month:0>2}-{days[1]:0>2}",
        freq="d"
    )
    dfs_sales.append(
        pd.DataFrame(
            index=weeks, data={
                "revenue": 2500 + 100*np.random.randn(len(weeks)),
                "customers": np.random.randint(10, 25)
            }
    ))

(jan_sales, feb_sales, mar_sales, apr_sales,
 may_sales, jun_sales, jul_sales, aug_sales,
 sep_sales, oct_sales, nov_sales, dec_sales) = dfs_sales
```

```python
stacked_sales = ...
stacked_sales
```

```python
grader.check("q9")
```

### Question 10:

Using the basketball dataset below, use `set_index`, `reset_index`, `stack`, and `unstack` to replicate the operations below:

```python
url = "https://datascience.quantecon.org/assets/data/bball.csv"
bball = pd.read_csv(url).query("Player != 'Ibaka'")
bball.info()

bball
```

```python
# Operation1 to replicate. call result `replicate_pivot`
bball.pivot_table(index="Year", columns="Player", values="Pts")
```

```python
# Operation2 to replicate. call result `replicate_melt`
bball.melt(id_vars=["Year", "Player", "Team"], value_vars=["Games", "Pts", "Assist"])
```

```python
replicate_melt = ...
replicate_pivot = ...
replicate_pivot
```

```python
grader.check("q10")
```

---

To double-check your work, the cell below will rerun all of the autograder tests.

```python
grader.check_all()
```

Tests for q1 (points: 5)
```
>>> assert df_full.shape == (4340, 3)
```
```
>>> assert df.shape == (217, 3)
```


Tests for q10 (points: 5)
```
>>> assert replicate_pivot.shape == (3, 2)
```
```
>>> assert set(replicate_pivot.columns) == {"Curry", "Durant"}
```
```
>>> assert replicate_melt.shape == (24, 5)
```
```
>>> assert set(replicate_melt.columns) == {"Year", "Player", "Team", "variable", "value"}
```


Tests for q2 (points: 5)
```
>>> assert df.shape == (217, 4)
```


Tests for q4 (points: 5)
```
>>> assert count.shape == (8, 10)
```
```
>>> assert len(xvals) == 9
```


Tests for q5 (points: 5)
```
>>> assert jd.shape == (8, 10)
```
```
>>> assert jd.sum() == 1
```


Tests for q6 (points: 5)
```
>>> assert len(marginal_pcr) == 10
```
```
>>> assert (marginal_pcr >= 0).all()
```
```
>>> assert len(marginal_lgdp) == 8
```
```
>>> assert (marginal_lgdp >= 0).all()
```


Tests for q7 (points: 5)
```
>>> assert conditionals_lgdp.shape == jd.shape
```
```
>>> assert np.allclose(conditionals_lgdp.sum(axis=0), 1)
```


Tests for q9 (points: 5)
```
>>> assert stacked_sales.shape == (366, 2)
```
